## Codio Activity 6.4: Adjusting Parameters for Variance

**Expected Time: 60 Minutes**

**Total Points: 20 Points**

This activity focuses on using the $\Sigma$ matrix to limit the principal components based on how much variance should be kept.  In the last activity, a scree plot was used to see when the difference in variance explained slows.  Here, you will determine how many components are required to explain a proportion of variance.  The dataset is a larger example of a housing dataset related to individual houses and features in Ames Iowa.  For our purposes the non-null numeric data is selected.

## Index:

- [Problem 1](#Problem-1)
- [Problem 2](#Problem-2)
- [Problem 3](#Problem-3)
- [Problem 4](#Problem-4)

```python
import numpy as np
from scipy.linalg import svd
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
from sklearn.datasets import fetch_openml
```

```python
#fetching the data
housing = fetch_openml(name="house_prices", as_frame=True)
```

```python
#examine the dataframe
housing.frame
```

Output:
```
Id  MSSubClass MSZoning  LotFrontage  LotArea Street Alley LotShape  \
0        1.0        60.0       RL         65.0   8450.0   Pave  None      Reg   
1        2.0        20.0       RL         80.0   9600.0   Pave  None      Reg   
2        3.0        60.0       RL         68.0  11250.0   Pave  None      IR1   
3        4.0        70.0       RL         60.0   9550.0   Pave  None      IR1   
4        5.0        60.0       RL         84.0  14260.0   Pave  None      IR1   
...      ...         ...      ...          ...      ...    ...   ...      ...   
1455  1456.0        60.0       RL         62.0   7917.0   Pave  None      Reg   
1456  1457.0        20.0       RL         85.0  13175.0   Pave  None      Reg   
1457  1458.0        70.0       RL         66.0   9042.0   Pave  None      Reg   
1458  1459.0        20.0       RL         68.0   9717.0   Pave  None      Reg   
1459  1460.0        20.0       RL         75.0   9937.0   Pave  None      Reg   

     LandContour Utilities  ... PoolA
```

```python
#select numeric data and drop missing values
df = housing.frame.select_dtypes(['float', 'int']).dropna(axis = 1)#.select_dtypes(['int', 'float'])
```

```python
df.shape
```

Output:
```
(1460, 35)
```

[Back to top](#Index:) 

## Problem 1

### Scale the data

**5 Points**

After selecting our numeric data, scale the data so that it is ready for SVD.  Assign the scaled data to `df_scaled` below.  Your answer should be of type DataFrame.

<!-- grade_id: cell-0aad0e887fead595, points: 0, grade: false, solution: true -->
```python
### GRADED

df_scaled = ''

### BEGIN SOLUTION
df_scaled = (df - df.mean())/df.std()
### END SOLUTION

# Answer check
print(type(df_scaled))
```

Output:
```
<class 'pandas.core.frame.DataFrame'>
```

<!-- grade_id: cell-ed93d779580e3c87, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
df_scaled_ = (df - df.mean())/df.std()
#
#
#
pd.testing.assert_frame_equal(df_scaled, df_scaled_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 2

### Extracting $\Sigma$

**5 Points**

Using the scaled data, extract the singular values from the data using the `scipy.linalg` function `svd`.  Assign your results to `sigma` below. 

<!-- grade_id: cell-c06f74cdd7fd8c26, points: 0, grade: false, solution: true -->
```python
### GRADED

U, sigma, VT = '', '', ''

### BEGIN SOLUTION
U, sigma, VT = svd(df_scaled)
### END SOLUTION

# Answer check
print(type(sigma))
print(sigma.shape)
```

Output:
```
<class 'numpy.ndarray'>
(35,)
```

<!-- grade_id: cell-102dc70bd31f2551, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
df_scaled_ = (df - df.mean())/df.std()
U_, sigma_, VT_ = svd(df_scaled_)
#
#
#
np.testing.assert_array_equal(sigma_, sigma)
assert sigma.shape == sigma_.shape
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 3

### Percent Variance Explained

**5 Points**

To compute the percent variance explained, we will divide each singular value by the sum of the singular values.  Assign your percents as an array to `percent_variance_explained` below.  Note that due to rounding this percent won't sum to exactly 1.  

<!-- grade_id: cell-d771d390de5194a6, points: 0, grade: false, solution: true -->
```python
### GRADED

percent_variance_explained = ''

### BEGIN SOLUTION
U, sigma, VT = svd(df_scaled)
percent_variance_explained = sigma/sigma.sum()
### END SOLUTION
print(percent_variance_explained.shape)
print(percent_variance_explained.sum())
```

Output:
```
(35,)
1.0000000000000002
```

<!-- grade_id: cell-f86a71c4d364e32f, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
df_scaled_ = (df - df.mean())/df.std()
U_, sigma_, VT_ = svd(df_scaled_)
percent_variance_explained_ = sigma_/sigma_.sum()
#
#
#
np.testing.assert_array_equal(percent_variance_explained, percent_variance_explained_)
### END HIDDEN TESTS
```

[Back to top](#Index:) 

## Problem 4

### Cumulative Variance Explained

**5 Points**

Using the solution to problem 3, how many principal components are necessary to retain 80% of the explained variance if we consider them in descending order?  Assign your response to `ans4` below as an integer. 

**HINT**: explore the `np.cumsum` function.

<!-- grade_id: cell-de2982e2498b5f70, points: 0, grade: false, solution: true -->
```python
### GRADED

ans4 = ''

### BEGIN SOLUTION
U, sigma, VT = svd(df_scaled)
percent_variance_explained = sigma/sigma.sum()
ans4 = int((np.cumsum(percent_variance_explained) < .8).sum())

### END SOLUTION
print(type(ans4))
print(ans4)
```

Output:
```
<class 'int'>
21
```

<!-- grade_id: cell-a57a0aada8255165, points: 5, grade: true, solution: false -->
```python
### BEGIN HIDDEN TESTS
df_scaled_ = (df - df.mean())/df.std()
U_, sigma_, VT_ = svd(df_scaled_)
percent_variance_explained_ = sigma_/sigma_.sum()
ans4_ = int((np.cumsum(percent_variance_explained_) < .8).sum())
#
#
#
assert ans4 == ans4_
### END HIDDEN TESTS
```